---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Introduction to Economic Modeling

```yaml
type: VideoExercise
key: bd0e6f6d3a
xp: 50
```

`@projector_key`
fcbbb89071cf5a91d502dded108df96b

---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise. I am doing it because the resources are not working and I am trying to figure out how to make a course here.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
